# MATH 303 Final Exam

## Student ID

Enter your student ID number in the cell below using the command:

```
id_student = XXXXXXXX # Replace XXXXXXXX with your student ID number
```

<!-- grade_id: cell-01e55b54e86dc9db, points: 0, grade: false, solution: true -->
```python
# Enter your student number
id_student =
```

Run the test in the cell below to verify your student ID number:

<!-- grade_id: cell-0b41f42dadd3932f, points: 0, grade: true, solution: false -->
```python
assert isinstance(id_student,int) , "Your student ID should be an integer."
print("Your student ID is",id_student)
```

## Problem 1

**The problem has two independent parts**:

* Part 1 (5 marks) includes questions a to e
* Part 2 (4 marks) includes questions f and g

## Part 1

We analyze a discrete time Markov chain with state space $\{1, 2, 3, 4\}$. Let $P$ be the transition matrix for this Markov chain.

Run the cell below to display $P$. The 4 states $1,2,3,4$ correspond to the row and column indices of $P$. To solve the following questions, **we recommend to draw the transition diagram associated with** $P$.

```python
from sympy import latex, Matrix
from IPython.display import display, Math
y = str(id_student); a = id_student + int(y[3]); bprob11 = a % 4;
Prob11 = [[[0.0,1.0,0.0,0.0],[0.4,0.0,0.6,0.0],[0.0,0.0,0.0,1.0],[0.0,0.0,0.5,0.5]],
          [[0.0,1.0,0.0,0.0],[0.0,0.0,1.0,0.0],[0.6,0.0,0.0,0.4],[0.0,0.0,0.0,1.0]],
          [[0.0,0.0,0.0,1.0],[0.0,0.0,0.6,0.4],[0.0,1.0,0.0,0.0],[0.0,0.0,0.0,1.0]],
          [[0.0,1.0,0.0,0.0],[0.5,0.0,0.5,0.0],[0.0,0.6,0.0,0.4],[0.0,0.0,0.0,1.0]]]
display(Math('P = ' + latex(Matrix(Prob11[bprob11]))))
```

### Problem 1a (1 mark)

Find the number of communicating classes `c` in this Markov chain. Enter your answer in the cell below.

**Grading:**

* +1 mark if answer `c` is correct
* -1 mark if answer `c` is incorrect

Alternatively, enter `c = 0` for 0 marks.

<!-- grade_id: cell-9c3642121a276e17, points: 0, grade: false, solution: true -->
```python
# Enter a value for c (or leave c = 0 for 0 marks)
c = 0
```

<!-- grade_id: cell-cf08eff50f982a06, points: 1, grade: true, solution: false -->
```python
"Verify the value c. This cell contains hidden tests."
assert isinstance(c,int) , "c should be an integer."
print("Problem 1a Test 1: Success!")
```

<!-- grade_id: cell-0da2e7ba186376ef, points: 1, grade: true, solution: false -->
```python
"Verify the value c. This cell contains hidden tests."
assert isinstance(c,int) , "c should be an integer."
print("Problem 1a Test 2: Success!")
```

### Problem 1b (1 mark)

Find all of the aperiodic states and enter them as a vector of integers (in increasing order) called `Ap`. For example, if you select states 1 and 4 then enter:
```
Ap = [1,4]
```
If you select only one state such as state 1 then enter:
```
Ap = [1]
```
If there are no aperiodic states, set `Ap` to the empty vector:
```
Ap = []
```

**Grading:**

* +1 mark if answer `Ap` is correct
* -1 mark if answer `Ap` is incorrect

Alternatively, enter `Ap = 0` for 0 marks.

<!-- grade_id: cell-031f5d2c78a406f6, points: 0, grade: false, solution: true -->
```python
# Enter a value for Ap (or leave Ap = 0 for 0 marks)
Ap = 0
```

<!-- grade_id: cell-b9b1be37f9db30e2, points: 1, grade: true, solution: false -->
```python
"Verify the value of Ap. This cell contains hidden tests."
assert isinstance(Ap,list) or Ap == 0, "Ap should be a list of integers or Ap = 0."
print("Problem 1b Test 1: Success!")
```

<!-- grade_id: cell-92004b917ba56858, points: 1, grade: true, solution: false -->
```python
"Verify the value of Ap. This cell contains hidden tests."
assert isinstance(Ap,list) or Ap == 0, "Ap should be a list of integers or Ap = 0."
print("Problem 1b Test 2: Success!")
```

### Problem 1c (1 mark)

Find all of the states of period 2 and enter them as a vector of integers (in increasing order) called `P2`. For example, if you select states 1 and 4 then enter:
```
P2 = [1,4]
```
If you select only one state such as state 1 then enter:
```
P2 = [1]
```
If there are no states of period 2, set `P2` to the empty vector:
```
P2 = []
```

**Grading:**

* +1 mark if answer `P2` is correct
* -1 mark if answer `P2` is incorrect

Alternatively, enter `P2 = 0` for 0 marks.

<!-- grade_id: cell-54bb36a1b34f8abc, points: 0, grade: false, solution: true -->
```python
# Enter a value for P2 (or leave P2 = 0 for 0 marks)
P2 = 0
```

<!-- grade_id: cell-3eea5ba2631c710c, points: 1, grade: true, solution: false -->
```python
"Verify the value of P2. This cell contains hidden tests."
assert isinstance(P2,list) or P2 == 0, "P2 should be a list of integers or P2 = 0."
print("Problem 1c Test 1: Success!")
```

<!-- grade_id: cell-dab628495512f40a, points: 1, grade: true, solution: false -->
```python
"Verify the value of P2. This cell contains hidden tests."
assert isinstance(P2,list) or P2 == 0, "P2 should be a list of integers or P2 = 0."
print("Problem 1c Test 2: Success!")
```

### Problem 1d (1 mark)

Find all of the states of period 3 and enter them as a vector of integers (in increasing order) called `P3`. For example, if you select states 1 and 4 then enter:
```
P3 = [1,4]
```
If you select only one state such as state 1 then enter:
```
P3 = [1]
```
If there are no states of period 3, set `P3` to the empty vector:
```
P3 = []
```

**Grading:**

* +1 mark if answer `P3` is correct
* -1 mark if answer `P3` is incorrect

Alternatively, enter `P3 = 0` for 0 marks.

<!-- grade_id: cell-62c2e533febc64cb, points: 0, grade: false, solution: true -->
```python
# Enter a value for P3 (or leave P3 = 0 for 0 marks)
P3 = 0
```

<!-- grade_id: cell-3f47d8a666542977, points: 1, grade: true, solution: false -->
```python
"Verify the value of P3. This cell contains hidden tests."
assert isinstance(P3,list) or P3 == 0, "P3 should be a list of integers or P3 = 0."
print("Problem 1d Test 1: Success!")
```

<!-- grade_id: cell-df3909c0b4365845, points: 1, grade: true, solution: false -->
```python
"Verify the value of P3. This cell contains hidden tests."
assert isinstance(P3,list) or P3 == 0, "P3 should be a list of integers or P3 = 0."
print("Problem 1d Test 2: Success!")
```

### Problem 1e (1 mark)

Find all of the absorbing states and enter them as a vector of integers (in increasing order) called `Ab`. For example, if you select states 1 and 4 then enter:
```
Ab = [1,4]
```
If you select only one state such as state 1 then enter:
```
Ab = [1]
```
If there are no absorbing states, set `Ab` to the empty vector:
```
Ab = []
```

**Grading:**

* +1 mark if answer `Ab` is correct
* -1 mark if answer `Ab` is incorrect

Alternatively, enter `Ab = 0` for 0 marks.

<!-- grade_id: cell-e04d45b766fc06cd, points: 0, grade: false, solution: true -->
```python
# Enter a value for Ab (or leave Ab = 0 for 0 marks)
Ab = 0
```

<!-- grade_id: cell-738a6adb2bdf342c, points: 1, grade: true, solution: false -->
```python
"Verify the value of Ab. This cell contains hidden tests."
assert isinstance(Ab,list) or Ab == 0, "Ab should be a list of integers or Ab = 0."
print("Problem 1e Test 1: Success!")
```

<!-- grade_id: cell-7196997212fa6db6, points: 1, grade: true, solution: false -->
```python
"Verify the value of Ab. This cell contains hidden tests."
assert isinstance(Ab,list) or Ab == 0, "Ab should be a list of integers or Ab = 0."
print("Problem 1e Test 2: Success!")
```

## Part 2

Let's consider a continuous time Markov chain $X(t)$ on $\{1, 2, 3, 4\}$. Run the following cell to get the initial state $x_0$, the transition probabilities $P=(p_{i,j})$, and the jump rates $V=(v_{i})$, as defined in the course. 

```python
from sympy import latex, Matrix
from IPython.display import display, Math

y = str(id_student); a = id_student + int(y[3]); bprob12 = a % 4;
Prob1fg = [[0.0,0.25,0.25,0.5],[0.4,0.0,0.2,0.4],[0.4,0,0.3,0.3],[0.5,0.0,0.5,0]]
x0 = [1,2,3,4]

display(Math("P = " + latex(Matrix(Prob1fg))))
display(Math("V = [2,3,5,4]"))
display(Math("x_0 = {}".format(x0[bprob12])))
```

### Problem 1f (2 marks)

Enter the value `r` in the cell below such that `t` samples the time of the first jump. See the [NumPy documentation](https://numpy.org/doc/1.16/reference/routines.random.html) if needed.

**Grading:**

* +2 marks if answer `r` is correct
* -1 mark if answer `r` is incorrect

Alternatively, enter `r = 0` for 0 marks.

<!-- grade_id: cell-fa56dc9bd364b344, points: 0, grade: false, solution: true -->
```python
import numpy as np

# Enter a value for r (or leave r = 0 for 0 marks)
r = 0

t = np.random.exponential(r)
```

<!-- grade_id: cell-969a070bed64a345, points: 1, grade: true, solution: false -->
```python
"Verify the value of r. This cell contains hidden tests."
assert r >= 0 , "r should be a nonnegative number."
print("Problem 1f Test 1: Success!")
```

<!-- grade_id: cell-37ca3d4a2a9edb9c, points: 2, grade: true, solution: false -->
```python
"Verify the value of r. This cell contains hidden tests."
assert r >= 0 , "r should be a nonnegative number."
print("Problem 1f Test 2: Success!")
```

### Problem 1g (2 marks)

Enter the value `Pi` in the cell below such that `x` samples the first arrival state (the state after the first jump). See the [NumPy documentation](https://numpy.org/doc/1.16/reference/routines.random.html) if needed.

**Grading:**

* +2 marks if answer `Pi` is correct
* -1 mark if answer `Pi` is incorrect

Alternatively, enter `Pi = 0` for 0 marks.

<!-- grade_id: cell-b548e824fe674d0c, points: 0, grade: false, solution: true -->
```python
import numpy as np

# Enter a value for Pi (or leave Pi = 0 for 0 marks)
Pi = 0

x = np.random.choice([1,2,3,4], 1, p=Pi)
```

<!-- grade_id: cell-f220321abdc80829, points: 1, grade: true, solution: false -->
```python
"Verify the value of Pi. This cell contains hidden tests."
assert Pi is not None
print("Problem 1g Test 1: Success!")
```

<!-- grade_id: cell-fabb89bf10221bda, points: 2, grade: true, solution: false -->
```python
"Verify the value of Pi. This cell contains hidden tests."
assert Pi is not None
print("Problem 1g Test 2: Success!")
```

## Problem 2

Run the following cell and note the values to be used in Problem 2.

```python
y = str(id_student); a = id_student + int(y[2]); b = a % 4;
Prob2 = [[6,3,2,3],[8,2,1,3],[8,4,1,4],[6,2,3,4]]
print("Transition probability P_{{RR}} = {0}/{1}".format(Prob2[b][2],Prob2[b][3]))
print("Average number of inches of snow in a snowy year =",Prob2[b][0])
print("Average number of inches of snow in a rainy year =",Prob2[b][1])
```

## Problem 3
Run the following cell and note the values to be found in Problem 3.

```python
y = str(id_student); a = id_student + int(y[4]); b = a % 3;
Prob3 = [["Probability that Jacob got at least one subscriber, and no followers who are not subscribers after a week.",
          "Probability  that  Jacob has exactly no subscriber and 2 followers who are not subscribers after 4 days.",
          "Probability that Jacob has at most 1 follower who is not subscriber and exactly 2 subscribers after 4 days."],
         ["Probability that exactly 3 followers are subscribers, given that Jacob got 8 total followers in two days.",
          "Probability that exactly 3 followers are not subscribers, given that Jacob got 6 total followers in 3 days.",
          "Probability that exactly 4 followers are subscribers, given that Jacob got 8 total followers in 4 days."],
         ["Expected number of followers who are not subscribers, given that Jacob got 8 total followers in two days.",
          "Expected number of subscribers, given that Jacob got 6 total followers in 3 days.",
          "Expected number of followers who are not subscribers, given that Jacob got 8 total followers in 4 days."]]
print(" Problem 3a:\n",Prob3[0][b],"\n")
print(" Problem 3b:\n",Prob3[1][b],"\n")
print(" Problem 3c:\n",Prob3[2][b],"\n")
```

## Problem 4
Run the following cell and note the values for $\mu_{\text{news}} =$ `mu_news`, $\mu_{\text{sports}} =$ `mu_sports`, $p$ and $p'$ to be used in Problem 4.

```python
y = str(id_student); a = 3*id_student + int(y[5]); b = a % 4;
Prob4 = [['1','1/2','1/2','1/3'],['1/2','1/3','1/3','1/2'],['1/2','1/3','1/4','2/3'],['1/4','1/2','1/3','1/4']]
print("mu_news =",Prob4[b][0])
print("mu_sports =",Prob4[b][1])
print("p =", Prob4[b][2])
print("p' =", Prob4[b][3])
```